```python
# Initialize Otter
import otter
grader = otter.Notebook("0000000000.ipynb")
```

# 1 ç”¨æˆ·å­˜æ¬¾è¡Œä¸ºé¢„æµ‹

### æ•°æ®é›†

é“¶è¡Œè¥é”€æ•°æ®é›†

#### èƒŒæ™¯ä»‹ç»

é‡‘èæœºæ„ä¸ºäº†åœ¨ä¸‹ä¸€æ¬¡è¥é”€æ´»åŠ¨å¯»æ‰¾æœ€ä½³çš„æ”¹è¿›ç­–ç•¥ï¼Œæ¢å¯»å¦‚ä½•åœ¨æœªæ¥çš„è¥é”€æ´»åŠ¨ä¸­å‘æŒ¥æ›´å¤§çš„æ•ˆåŠ›ã€‚ä¸ºäº†è§£å†³è¿™ä¸ªé—®é¢˜ï¼Œæ•°æ®åˆ†æå¸ˆå¿…é¡»åˆ†æè¯¥é“¶è¡Œæœ€è¿‘å¼€å±•çš„è¥é”€æ´»åŠ¨ï¼Œé€šè¿‡ä¹‹å‰æ´»åŠ¨çš„æƒ…å†µé¢„æµ‹ç”¨æˆ·æ˜¯å¦ä¼šå‚åŠ ä¹‹åçš„è¥é”€æ´»åŠ¨ã€‚

#### ç‰¹å¾ä»‹ç»

##### ä¸ªäººä¿¡æ¯ç›¸å…³
1 - age : å¹´é¾„ (æ•°å€¼åˆ—)

2 - job : èŒä¸šçš„ç§ç±»(ç±»åˆ«åˆ—: 'admin.','blue-collar','entrepreneur','housemaid','management','retired','self-employed','services','student','technician','unemployed','unknown')

3 - marital : å©šå§»çŠ¶å†µ (ç±»åˆ«åˆ—: 'divorced','married','single','unknown')

4 - education : å—æ•™è‚²æƒ…å†µ(ç±»åˆ«åˆ—: 
'basic.4y','basic.6y','basic.9y','high.school','illiterate','professional.course','university.degree','unknown')

5 - default: æ˜¯å¦ä¿¡ç”¨è¿çº¦ (ç±»åˆ«åˆ—: 'no','yes','unknown')

6 - balance: ä½™é¢ (æ•°å€¼åˆ—)

7 - housing: æ˜¯å¦æœ‰ä½æˆ¿è´·æ¬¾(ç±»åˆ«åˆ—: 'no','yes','unknown')

8 - loan: æ˜¯å¦æœ‰ä¸ªäººè´·æ¬¾? (ç±»åˆ«åˆ—: 'no','yes','unknown')



##### ä¸å½“å‰è¥é”€æ´»åŠ¨çš„æœ€åä¸€æ¬¡è”ç³»ï¼š
9 - contact: è”ç³»äººé€šä¿¡ç±»å‹ (ç±»åˆ«åˆ—: 'cellular','telephone')

10 - month: æœ€åä¸€æ¬¡è”ç³»çš„æœˆä»½ (ç±»åˆ«åˆ—: 'jan', 'feb', 'mar', ..., 'nov', 'dec')

11 - day: æœ€åä¸€æ¬¡è”ç³»åœ¨æ˜ŸæœŸå‡ ï¼Ÿ (ç±»åˆ«åˆ—: 'mon','tue','wed','thu','fri')

12 - duration: æœ€åä¸€æ¬¡è”ç³»çš„æ—¶é•¿, å•ä½ç§’ (æ•°å€¼åˆ—). 

##### å…¶ä»–çš„ä¸€äº›ç‰¹å¾ï¼š
13 - campaign: æ­¤æ´»åŠ¨æœŸé—´å’Œæ­¤å®¢æˆ·çš„è”ç³»äººæ•° (æ•°å€¼åˆ—)

14 - pdays: ä¸Šæ¬¡æ´»åŠ¨ä¸­è”ç³»å®¢æˆ·åç»è¿‡çš„å¤©æ•° (æ•°å€¼åˆ—; 999 è¡¨ç¤ºè¿‘æœŸæ— è”ç³»)

15 - previous: æ­¤æ´»åŠ¨ä¹‹å‰å’Œæ­¤å®¢æˆ·ç«¯æ‰§è¡Œçš„è”ç³»äººæ•° (æ•°å€¼åˆ—)

16 - poutcome: ä¸Šä¸€æ¬¡è¥é”€æ´»åŠ¨çš„ç»“æœ (ç±»åˆ«åˆ—: 'failure','nonexistent','success')




### ç›®æ ‡

##### é¢„æµ‹å®šæœŸå­˜æ¬¾é¢åº¦

17 - deposit: æ˜¯å¦ä¼šå­˜æ¬¾ (ç±»åˆ«åˆ—: 'yes','no')


- æ•°æ®æ¥æºï¼šhttps://www.kaggle.com/datasets/janiobachmann/bank-marketing-dataset

## 1.1 ç¯å¢ƒå¯¼å…¥å’Œæ•°æ®å‡†å¤‡

```python
import numpy as np
import pandas as pd
from sklearn.preprocessing import LabelEncoder
from sklearn.metrics import accuracy_score
from sklearn.model_selection import train_test_split
from sklearn.svm import SVC 
from sklearn.linear_model import LogisticRegression
rng_seed = 44
```

```python
df=pd.read_csv("bank.csv")
df.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 11162 entries, 0 to 11161
Data columns (total 17 columns):
 #   Column     Non-Null Count  Dtype 
---  ------     --------------  ----- 
 0   age        11162 non-null  int64 
 1   job        11162 non-null  object
 2   marital    11162 non-null  object
 3   education  11162 non-null  object
 4   default    11162 non-null  object
 5   balance    11162 non-null  int64 
 6   housing    11162 non-null  object
 7   loan       11162 non-null  object
 8   contact    11162 non-null  object
 9   day        11162 non-null  int64 
 10  month      11162 non-null  object
 11  duration   11162 non-null  int64 
 12  campaign   11162 non-null  int64 
 13  pdays      11162 non-null  int64 
 14  previous   11162 non-null  int64 
 15  poutcome   11162 non-null  object
 16  deposit    11162 non-null  object
dtypes: int64(7), object(10)
memory usage: 1.4+ MB
```

```python
from sklearn.preprocessing import LabelEncoder
x = df.iloc[:,:-1]
y = df.iloc[:,-1]

print(len(x))
x_simple_prepared = pd.get_dummies(x.copy())

lbe =  LabelEncoder()
lbe.fit(y)
y = lbe.transform(y)
x_simple_prepared
```

Output:
```
11162
       age  balance  day  duration  campaign  pdays  previous  job_admin.  \
0       59     2343    5      1042         1     -1         0           1   
1       56       45    5      1467         1     -1         0           1   
2       41     1270    5      1389         1     -1         0           0   
3       55     2476    5       579         1     -1         0           0   
4       54      184    5       673         2     -1         0           1   
...    ...      ...  ...       ...       ...    ...       ...         ...   
11157   33        1   20       257         1     -1         0           0   
11158   39      733   16        83         4     -1         0           0   
11159   32       29   19       156         2     -1         0           0   
11160   43        0    8         9         2    172         5           0   
11161   34        0    9       628         1     -1         0           0   

       job_blue-collar  job_entrepreneur  ...  month_jun  month_mar  
```

## 1.2 æ¨¡å‹çš„è®­ç»ƒå’Œè¯„ä¼°(20åˆ†)

_Type your answer here, replacing this text._

```python
np.random.seed(rng_seed)
# åˆ©ç”¨get_dummies()æ–¹æ³•å¯¹xè¿›è¡Œç¼–ç 
x = df.iloc[:,:-1]
y = df.iloc[:,-1]
x_simple_prepared = pd.get_dummies(x.copy())
# åˆ©ç”¨LabelEncoder()æ–¹æ³•å¯¹yè¿›è¡Œç¼–ç 
lbe =  LabelEncoder()
lbe.fit(y)
y = lbe.transform(y)

np.random.seed(rng_seed)
# åˆ’åˆ†train, testæ•°æ®é›†
x_train, x_test, y_train, y_test = train_test_split(x_simple_prepared, y, train_size=0.8, test_size=1-0.8, random_state=0)
#è®­ç»ƒSVCæ¨¡å‹
svc_model = SVC(random_state=0, probability=True)
# è¿›è¡Œé¢„æµ‹
svc_simple_score = accuracy_score(y_test, svc_model.fit(x_train, y_train).predict(x_test))
# è®­ç»ƒ LR æ¨¡å‹
lr_model = LogisticRegression(random_state=0, max_iter=10000)
# è¿›è¡Œé¢„æµ‹
lr_simple_score = accuracy_score(y_test, lr_model.fit(x_train, y_train).predict(x_test))

print('svc accuracy:', svc_simple_score)
print('lr accuracy:', lr_simple_score)
```

Output:
```
svc accuracy: 0.7357814599193909
lr accuracy: 0.8248992386923422
```

```python
grader.check("p1_1")
```

Output:
```
p1_1 results: All test cases passed!
```

## 1.3 æ•°æ®é¢„å¤„ç†å’Œå†è®­ç»ƒ(20åˆ†)

**æ•°æ®é¢„å¤„ç†**å¯ä»¥å‚è€ƒçš„æ•°æ®å‡†å¤‡æ“ä½œï¼šhttps://scikit-learn.org/stable/modules/preprocessing.html#preprocessing
* é’ˆå¯¹`x_train`, `y_train`, `x_test`, `y_test`ä½¿ç”¨æ•°æ®å‡†å¤‡æ“ä½œï¼Œä½†è¦ä¿è¯
  * ä¸èƒ½å¢åŠ æˆ–åˆ å‡`x_test`å’Œ`y_test`çš„è¡Œï¼Œä½†å¯ä»¥å¯¹å®ƒä»¬è¿›è¡Œå˜æ¢
  * å¯ä»¥å¯¹`x_train`å’Œ`y_train`åšä»»ä½•æ“ä½œ

```python

# æ•°æ®é¢„å¤„ç†
from sklearn.preprocessing import StandardScaler
from sklearn.impute import SimpleImputer
from sklearn.decomposition import PCA
# å¤„ç†å¼‚å¸¸å€¼

threshold=3
q1 = np.percentile(x_train, 25, axis=0)
q3 = np.percentile(x_test, 75, axis=0)
print(x_train.head())
# è®¡ç®—IQRï¼ˆå››åˆ†ä½æ•°èŒƒå›´ï¼‰
iqr = q3 - q1
# è®¡ç®—ä¸Šä¸‹ç•Œ
lower_bound = q1 - threshold * iqr
upper_bound = q3 + threshold * iqr
# å°†å¼‚å¸¸å€¼æ›¿æ¢ä¸ºç¼ºå¤±å€¼ï¼Œå¹¶è¿›è¡Œæ’è¡¥
is_outlier = (x_train < lower_bound) | (x_train > upper_bound)
x_train_no_outlier = x_train.copy()
x_train_no_outlier[is_outlier] = np.nan
x_train_no_outlier = SimpleImputer(strategy='mean').fit_transform(x_train_no_outlier)
x_test_no_outlier = x_test.copy()
x_test_no_outlier[is_outlier] = np.nan
x_test_no_outlier = SimpleImputer(strategy='mean').fit_transform(x_test_no_outlier)

# åˆ†ç¦»ç‰¹å¾å’Œæ ‡ç­¾ï¼Œå¯¹æ ‡ç­¾åšç¼–ç 
x_train_no_outlier = pd.DataFrame(x_train_no_outlier )
x_test_no_outlier = pd.DataFrame(x_test_no_outlier)

# å¯¹ç‰¹å¾åšç¼–ç 
x_train_no_outlier = pd.get_dummies(x_train_no_outlier)
x_test_no_outlier = pd.get_dummies(x_test_no_outlier)
# å¯¹ç‰¹å¾åšé™ç»´
pca = PCA(n_components=20)
pca.fit(x_train_no_outlier)
x_train_no_outlier = pca.transform(x_train_no_outlier)
x_test_no_outlier = pca.transform(x_test_no_outlier)
# å¯¹ç‰¹å¾åšå½’ä¸€åŒ–
scaler = StandardScaler()
scaler.fit(x_train_no_outlier)
x_train_no_outlier = scaler.transform(x_train_no_outlier)
x_test_no_outlier = scaler.transform(x_test_no_outlier)
#å¯¹ç‰¹å¾åšç‰¹å¾å·¥ç¨‹
# ç‰¹å¾å·¥ç¨‹
# from sklearn.preprocessing import PolynomialFeatures
# poly = PolynomialFeatures(degree=2)
# x_train_no_outlier = poly.fit_transform(x_train_no_outlier)
# x_test_no_outlier = poly.transform(x_test_no_outlier)
np.random.seed(rng_seed)
# è®­ç»ƒSVCæ¨¡å‹ï¼Œå¤„ç†å®Œæ•°æ®åï¼Œå†æ¬¡è®­ç»ƒï¼Œä½ éœ€è¦å¾—åˆ°ä¸€ä¸ªsvc_scoreçš„æŒ‡æ ‡
svc_new_model=SVC(random_state=0, probability=True)
svc_score=accuracy_score(y_test, svc_new_model.fit(x_train_no_outlier, y_train).predict(x_test_no_outlier))
#y_prob = svc_new_model.predict_proba(x_test)
#è®­ç»ƒLRæ¨¡å‹
lr_new_model=LogisticRegression(random_state=0, max_iter=1000)
lr_score=accuracy_score(y_test, lr_new_model.fit(x_train_no_outlier, y_train).predict(x_test_no_outlier))
#y_prob1=lr_new_model.predict_proba(x_test)
print('svc accuracy:', svc_score)
print('lr accuracy:', lr_score)
```

Output:
```
age  balance  day  duration  campaign  pdays  previous  job_admin.  \
6111   49      104    8       285         1     -1         0           0   
1149   59      514    3       261         1     -1         0           0   
2954   33      195   16       984         5     -1         0           0   
2458   58      687    7       221         1     99         2           0   
7504   33     1077   18       432         3    374         3           1   

      job_blue-collar  job_entrepreneur  ...  month_jun  month_mar  month_may  \
6111                0                 0  ...          0          0          1   
1149                0                 0  ...          0          0          0   
2954                1                 0  ...          0          0          0   
2458                0                 0  ...          0          0          0   
7504                0                 0  ...          0          0          1   

      month_nov  month_oct  month_sep  poutcome_failure  poutc
```

```python
grader.check("p1_2")
```

Output:
```
p1_2 results: All test cases passed!
```

# 2 é€»è¾‘å›å½’åˆ†ç±»å™¨çš„å®ç°

## 2.1 ç¡®å®šä¼˜åŒ–ç›®æ ‡(10åˆ†)

æˆ‘ä»¬é‡‡ç”¨äºŒå…ƒäº¤å‰ç†µä½œä¸ºæŸå¤±å‡½æ•°ï¼Œå…·ä½“çš„æ¨ç†è¿‡ç¨‹å¯å‚è€ƒä¸»è¯¾[Lec09](../../Lectures/Lec09/)

$J(\theta) = -\frac{1}{m} \sum_{i=1}^{m} \left[ y^{(i)} \log(\hat{y}^{(i)}) + (1 - y^{(i)}) \log(1 - \hat{y}^{(i)}) \right]$

* $J(\theta)$æ˜¯æŸå¤±å‡½æ•°
* $m$æ˜¯æ ·æœ¬æ•°é‡
* $y^{(i)}$æ˜¯ç¬¬$i$ä¸ªæ ·æœ¬çš„çœŸå®æ ‡ç­¾
* $\hat{y}^{(i)}$æ˜¯é¢„æµ‹æ¦‚ç‡
* $\theta$æ˜¯æƒé‡å‘é‡

```python
np.random.seed(rng_seed)
# å®ç°æŸå¤±å‡½æ•°
def sigmoid(z):
    return 1 / (1 + np.exp(-z))

def logistic_loss(y, y_hat):
    epsilon = 1e-15  # ä¸€ä¸ªå¾ˆå°çš„ epsilon å€¼ï¼Œç”¨äºé¿å…å–é›¶çš„å¯¹æ•°
    #print(y_hat)
    logistic_loss = -np.sum(y * np.log(y_hat) + (1 - y) * np.log(1 - y_hat))
    logistic_loss = logistic_loss / len(y)
    
    return logistic_loss

def loss_function(X, y, weights):
    

    z = np.dot(X, weights)
    y_hat = sigmoid(z)
    loss = logistic_loss(y, y_hat)
    return loss
```

```python
grader.check("p2_1")
```

Output:
```
p2_1 results: All test cases passed!
```

## 2.2 è®¡ç®—ä¼˜åŒ–ç›®æ ‡çš„æ¢¯åº¦(10åˆ†)
è¯·æ ¹æ®ä¸Šè¿°ç›®æ ‡è®¡ç®—æ¢¯åº¦

$\nabla_{\theta} J(\theta) = \frac{1}{m} X^T (\hat{y} - y)$

è¯·æ ¹æ®ç»™å‡ºçš„å…¬å¼ï¼Œç¼–å†™æ¢¯åº¦å‡½æ•°ã€‚æ³¨æ„ $\hat{y}$ è¦é€šè¿‡sigmoidå‡½æ•°åä¸ $y$ ç›¸å‡

```python
np.random.seed(rng_seed)
# å®ç°æ¢¯åº¦å‡½æ•°
def gradient(X, y, weights):
    z = np.dot(X, weights)
    y_hat = sigmoid(z)
    gradient = np.dot(X.T, y_hat - y) / X.shape[0]
    return gradient
```

```python
grader.check("p2_2")
```

Output:
```
p2_2 results: All test cases passed!
```

## 2.3 ä½¿ç”¨éšæœºæ¢¯åº¦ä¸‹é™ï¼ˆSGDï¼‰è¿›è¡Œä¼˜åŒ–(10åˆ†)

æƒé‡æ›´æ–°çš„æ•°å­¦è¡¨ç¤ºä¸ºï¼š

$\theta = \theta - \alpha \nabla_{\theta} J(\theta)$

```python
def stochastic_gradient_descent(X, y, weights, learning_rate=0.01, num_iterations=100):
    m = y.size
    loss_history = []
    
    for i in range(num_iterations):
        for j in range(m):
            # éšæœºé€‰æ‹©ä¸€ä¸ªæ•°æ®ç‚¹
            idx = np.random.randint(m)
            X_i = X[idx, :].reshape(1, -1)
            y_i = y[idx]
            # è®¡ç®—æ¢¯åº¦å¹¶æ›´æ–°æƒé‡
            grad =gradient(X_i, y_i, weights)
            weights -= learning_rate * grad
            
            # è®¡ç®—å¹¶è®°å½•æŸå¤±ï¼Œç”¨äºç›‘æ§
            loss = loss_function(X, y, weights)
            if(loss==np.nan):
                print(i,j)
            loss_history.append(loss)
        
        # æ¯10æ¬¡è¿­ä»£æ‰“å°ä¸€æ¬¡æŸå¤±å€¼
        if (i+1) % 1 == 0:
            print(f"Iteration {i+1}/{num_iterations}, Loss: {loss:.4f}")
    
    return weights, loss_history

np.random.seed(rng_seed)
num_features = 6
num_samples = 138
dummy_weights = np.random.rand(num_features)
dummy_X = np.random.rand(num_samples, num_features)
dummy_y = np.random.randint(0, 2, num_samples)

initial_weights = np.zeros(num_features)

final_weights, loss_history = stochastic_gradient_descent(dummy_X, dummy_y, initial_weights)
```

Output:
```
Iteration 1/100, Loss: 0.6920
Iteration 2/100, Loss: 0.6907
Iteration 3/100, Loss: 0.6898
Iteration 4/100, Loss: 0.6904
Iteration 5/100, Loss: 0.6894
Iteration 6/100, Loss: 0.6887
Iteration 7/100, Loss: 0.6879
Iteration 8/100, Loss: 0.6870
Iteration 9/100, Loss: 0.6860
Iteration 10/100, Loss: 0.6853
Iteration 11/100, Loss: 0.6849
Iteration 12/100, Loss: 0.6847
Iteration 13/100, Loss: 0.6847
Iteration 14/100, Loss: 0.6844
Iteration 15/100, Loss: 0.6845
Iteration 16/100, Loss: 0.6841
Iteration 17/100, Loss: 0.6839
Iteration 18/100, Loss: 0.6835
Iteration 19/100, Loss: 0.6832
Iteration 20/100, Loss: 0.6831
Iteration 21/100, Loss: 0.6836
Iteration 22/100, Loss: 0.6823
Iteration 23/100, Loss: 0.6846
Iteration 24/100, Loss: 0.6823
Iteration 25/100, Loss: 0.6828
Iteration 26/100, Loss: 0.6811
Iteration 27/100, Loss: 0.6814
Iteration 28/100, Loss: 0.6816
Iteration 29/100, Loss: 0.6813
Iteration 30/100, Loss: 0.6817
Iteration 31/100, Loss: 0.6817
Iteration 32/100, Loss: 0.6806
Iteration 33/100,
```

```python
grader.check("p2_3")
```

Output:
```
p2_3 results: All test cases passed!
```

## 2.4 è®­ç»ƒä¸è¯„æµ‹(10åˆ†)

* ä½¿ç”¨ä¸Šé¢å®ç°çš„å‡½æ•°è®­ç»ƒä¸€ä¸ªLRåˆ†ç±»å™¨
* ä¸ºäº†ä¿è¯ç»“æœæ–¹ä¾¿æµ‹è¯•...

```python
x_train, x_test, y_train, y_test = train_test_split(x_simple_prepared, y, train_size=0.8, test_size=1-0.8, random_state=0)
x_train = x_train.values
x_test = x_test.values
```

```python
from sklearn.metrics import classification_report
def predict(X, weights):
    # æ ¹æ®ç»™å®šçš„æƒé‡ï¼Œé¢„æµ‹æ•°æ®é›†Xçš„æ ‡ç­¾
    z = np.dot(X, weights)
    z=sigmoid(z)
    predictions = (z >= 0.5).astype(int)
    return predictions  # é¢„æµ‹æ¦‚ç‡å¤§äºç­‰äº0.5è§†ä¸ºç±»åˆ«1ï¼Œå¦åˆ™ä¸ºç±»åˆ«0
    
num_features = x_train.shape[1]
np.random.seed(rng_seed)
# æ ‡å‡†åŒ–x_trainå’Œx_testï¼Œå…¬å¼å¦‚ä¸‹
# x = (x-mean(x))/std(x)
x_train=(x_train-np.mean(x_train,axis=0))/(np.std(x_train,axis=0))
x_test=(x_test-np.mean(x_test,axis=0))/np.std(x_test,axis=0)
# weightåˆå§‹åŒ–ä¸º0
weights = np.zeros(num_features)
learning_rate = 0.01
num_iterations = 10
final_weights, loss_history = stochastic_gradient_descent(x_train, y_train, weights, learning_rate, num_iterations)

predictions = predict(x_test, final_weights)

accuracy = accuracy_score(y_test, predictions)
report = classification_report(y_test, predictions)
```

Output:
```
Iteration 1/10, Loss: 0.4242
Iteration 2/10, Loss: 0.4219
Iteration 3/10, Loss: 0.4258
```

```python
predictions = predict(x_test, final_weights)
accuracy = accuracy_score(y_test, predictions)
report = classification_report(y_test, predictions)
loss_history
print(report)
print(accuracy)
```

Output:
```
precision    recall  f1-score   support

           0       0.82      0.84      0.83      1205
           1       0.80      0.78      0.79      1028

    accuracy                           0.81      2233
   macro avg       0.81      0.81      0.81      2233
weighted avg       0.81      0.81      0.81      2233

0.8101209135691895
```

```python
grader.check("p2_4")
```

Output:
```
0.691986260605109
```

# 3 åŸºäºå®éªŒç»“æœçš„åˆ†æä¸ç ”è®¨

è¯·å¯¹æœ¬æ¬¡ä½œä¸šåšå‡ºæ€»ç»“ï¼š

- (1) ä½ é‡‡ç”¨çš„**æ•°æ®é¢„å¤„ç†æ“ä½œ**æ˜¯å¦æå‡äº†æ¨¡å‹çš„accuracy(5åˆ†)?å¦‚æœæœ‰æé«˜ï¼Œè¯·ç»™å‡ºæé«˜çš„å…·ä½“æ•°å€¼ï¼Œå¹¶åˆ†ææå‡çš„åŸå› (15åˆ†)ã€‚
- (2) é’ˆå¯¹å¤ç°çš„é€»è¾‘å›å½’æ¨¡å‹ï¼Œå®ƒå’Œç›´æ¥ä½¿ç”¨sklearnçš„ç²¾åº¦å·®è·ä¸ºå¤šå¤§(10åˆ†)ï¼Ÿè¯·åˆ†æå…¶èƒŒåçš„åŸå› (10åˆ†)ã€‚

<font color='green'>**ç­”:**</font>
æ˜¯çš„ï¼Œå¯¹äºSVMæ¨¡å‹ï¼Œæé«˜äº†0.08,å¯¹äºLRæ¨¡å‹ï¼Œæé«˜äº†0.01  
æå‡çš„å¯èƒ½åŸå› ï¼š
* æ ‡å‡†åŒ–ä¹‹åå»é™¤äº†å¼‚å¸¸å€¼ï¼Œé™ä½äº†å¼‚å¸¸å€¼çš„å½±å“ã€‚
* ä¸ºäº†åœ¨

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Output:
```
File d:\anaconda\envs\newfortry\lib\site-packages\otter\export\__init__.py:36, in export_notebook(nb_path, dest, exporter_type, **kwargs)
     33     pdf_name = os.path.splitext(nb_path)[0] + ".pdf"
     35 Exporter = get_exporter(exporter_type=exporter_type)
---> 36 Exporter.convert_notebook(nb_path, pdf_name, **kwargs)
     38 return pdf_name

File d:\anaconda\envs\newfortry\lib\site-packages\otter\export\exporters\via_latex.py:70, in PDFViaLatexExporter.convert_notebook(cls, nb_path, dest, xecjk, **kwargs)
     67     with open(dest, "wb") as output_file:
     68         output_file.write(pdf_output[0])
---> 70 except nbconvert.pdf.LatexFailed as error:
     71     message = "There was an error generating your LaTeX; showing full error message:\n"
     72     message += indent(error.output, "    ")

AttributeError: module 'nbconvert' has no attribute 'pdf'
```

Tests for p1_1 (points: 20)
```
>>> df['age'][2] == 41 and df['balance'][48] == 1443
True
```
```
>>> abs(svc_simple_score - 0.7357814599193909) < 0.01
True
```


Tests for p1_2 (points: 20)
```
>>> svc_score - svc_simple_score>=0
True
```
```
>>> lr_score - lr_simple_score>=0
True
```


Tests for p2_1 (points: 10)
```
>>> num_features = 5
>>> num_samples = 10
>>> np.random.seed(rng_seed)
>>> dummy_weights = np.random.rand(num_features)
>>> dummy_X = np.random.rand(num_samples, num_features)
>>> dummy_y = np.random.randint(0, 2, num_samples)
>>> 
>>> initial_loss = loss_function(dummy_X, dummy_y, dummy_weights)
>>> abs(initial_loss - 0.9160538386731103) < 0.0001
True
```


Tests for p2_2 (points: 10)
```
>>> num_features = 5
>>> num_samples = 10
>>> np.random.seed(rng_seed)
>>> dummy_weights = np.random.rand(num_features)
>>> dummy_X = np.random.rand(num_samples, num_features)
>>> dummy_y = np.random.randint(0, 2, num_samples)
>>> 
>>> initial_gradient = gradient(dummy_X, dummy_y, dummy_weights)
>>> abs(initial_gradient.sum() - 0.7627876529375204) < 0.0001
True
```


Tests for p2_3 (points: 10)
```
>>> abs(loss_history[0]-0.6933603914753861) < 0.0001
True
```
```
>>> abs(loss_history[1]-0.6931379470500929) < 0.0001
True
```
```
>>> abs(len(loss_history)-13800) < 0.0001
True
```


Tests for p2_4 (points: 10)
```
>>> x_train.shape == (8929, 51)
True
```
```
>>> abs(float(final_weights[0]) - (-0.041616572346357336)) < 0.0001
True
```
```
>>> abs(loss_history[0] - 0.6898603068851875) < 0.0001
True
```
